# Assignment 3

#### Submission Guidelines:
1. Save and run all the cells of this notebook before submitting.
2. Use the Validate button located on the jupyter notebook toolbar to check your solution. If the solution passes all the tests, a success pop-up will show. Otherwise, a pop-up will show which cells have failed.
4. You will be able to test your work on the local test cases by running the cell with assertion commands.
5. There are "Hidden Test Blocks" which will test the solutions on hidden test cases. Make sure to run all the blocks to ensure the test cases are tested.
6. Once you are done with your code, you should save the notebook and submit it in the assignment section on datahub. Ensure that you submit your assignment using the "Submit" button on the Assignment section.
7. To download the file as an .py, navigate to File -> Download as -> .py
8. Plagiarism and AI similarity check will be conducted.

In this assignment, you will create a class to manage a multivariate polynomial. E.g. $P(x_1, x_2) = 1 - x_1 - x_2 + x_1 x_2$.

A multivariate polynomial is mathematically expressed as:

$$P(x_1, \dots, x_n) \triangleq \sum_{i_1 = 0}^{d_1} \cdots \sum_{i_n = 0}^{d_n} p_{i_1, \dots, i_n} x_1^{i_1} \cdots x_n^{i_n}$$

, where $d_j$ is the maximum degree of variable $x_j$ and $p_{i_1, \dots, i_n}$ are the coefficients corresponding to each term in the polynomial.

The following questions will guide you through the process from defining the polynomial structure to implementing basic functionality.

```python
# DO NOT MODIFY THE FOLLOWING LINE(S)
import numpy as np
import xarray as xr
from typing import List, Tuple, Dict, Optional
```

## 1: Define the Multivariate Polynomial Class (25 points)

Create a `MultivariatePolynomial` class that uses `xarray.DataArray` to store the coefficients of a multivariate polynomial. The class should allow for any number of variables and positive integer degrees.

### 1.1 Initialisation (5 points)

Under the definition of the class, implement the `__init__` method that takes the following parameters:
- `variables`: A list of variable names (strings). E.g., `['x_1', 'x_2']`.
- `degrees`: A list of maximum degrees for each variable (positive integers). E.g., `[1, 1]`.
- `coefficients`: A dictionary containing the coefficients of the polynomial, where keys are tuples representing the degrees of each variable, and values are the corresponding coefficients. E.g., `{(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}`.

Notes:
- You should create an `xarray.DataArray` to store the coefficients.
- The shape of the DataArray should be determined by the maximum degrees of each variable.
- The coordinates of the DataArray should represent the degrees of each variable.

<!-- grade_id: cell-1d5928974f1349ba, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### 1.2 Get Coefficient (10 points)

Implement a method `get_coefficient(self, degree_tuple)` that takes a tuple representing the degrees of each variable and returns the corresponding coefficient from the DataArray.

Return `None` if the degree tuple is not valid (either the tuple length does not match the number of variables or any degree exceeds the maximum degree for that variable).

<!-- grade_id: cell-f81b56fdf93664ca, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4a9a9bfa40f7aefc, points: 2, grade: true, solution: false -->
```python
# Test case 1.
variables = ['x_1', 'x_2']
degrees = [1, 1]
coefficients = {(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}
poly = MultivariatePolynomial(variables, degrees, coefficients)
assert poly.get_coefficient((1, 1)) == 1
```

<!-- grade_id: cell-f3f984de81777f78, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-3a15f3f5cd2d491e, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2.
```

### 1.3 Set Coefficient (10 points)
Implement a method `set_coefficient(self, degree_tuple, value)` that takes a tuple representing the degrees of each variable and a value to set the corresponding coefficient in the DataArray.

Do nothing if the degree tuple is not valid (either the tuple length does not match the number of variables or any degree exceeds the maximum degree for that variable).

Note:
- The method `get_coefficient` will be used to verify the correctness of this method in the test cases. Make sure it works correctly before proceeding.

<!-- grade_id: cell-34432b871e2bbd5f, points: 4, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-dbfe4839ae7c8270, points: 2, grade: true, solution: false -->
```python
# Test case 1.
variables = ['x_1', 'x_2']
degrees = [1, 1]
coefficients = {(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}
poly = MultivariatePolynomial(variables, degrees, coefficients)
poly.set_coefficient((1, 1), -1)
assert poly.get_coefficient((1, 1)) == -1
```

<!-- grade_id: cell-44b3632ab3cf9916, points: 2, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-4a1a8364b8d7e332, points: 2, grade: true, solution: false -->
```python
# Hidden test case 2.
```

## 2. Polynomial Evaluation (25 points)

Implement a method `evaluate(self, values)` that takes a tuple of variable values and evaluates the polynomial at those values.

Return `None` if the length of the values tuple does not match the number of variables.

Hint: You may use `numpy.ndindex` to iterate over all possible degree combinations. For more help please visit https://numpy.org/doc/2.3/reference/generated/numpy.ndindex.html

<!-- grade_id: cell-3e7e6e5019191c20, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-ffa93233536136f6, points: 5, grade: true, solution: false -->
```python
# Test case 1.
variables = ['x_1', 'x_2']
degrees = [1, 1]
coefficients = {(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}
poly = MultivariatePolynomial(variables, degrees, coefficients)
assert poly.evaluate((1, 1)) == 0
```

<!-- grade_id: cell-69b1ac7c159e2232, points: 5, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-b42ba849ff3b12d9, points: 5, grade: true, solution: false -->
```python
# Hidden test case 2.
```

<!-- grade_id: cell-521fbac01a48cb5f, points: 5, grade: true, solution: false -->
```python
# Hidden test case 3.
```

# 3. Polynomial Addition (25 points)
Implement a method `add(self, other)` that takes another `MultivariatePolynomial` object and returns a new `MultivariatePolynomial` object representing the sum of the two polynomials.

The mathematical expression for the addition is:

$$P(x_1, \dots, x_n) + Q(x_1, \dots, x_n) = \sum_{i_1 = 0}^{d_1} \cdots \sum_{i_n = 0}^{d_n} (p_{i_1, \dots, i_n} + q_{i_1, \dots, i_n}) x_1^{i_1} \cdots x_n^{i_n}$$

For simplicity, we assume that both polynomials have the same variables and degrees. You do not need to handle cases where they differ. Return `None` if they do not match.

Hint: You may use `numpy.ndindex` to iterate over all possible degree combinations. For more help please visit https://numpy.org/doc/2.3/reference/generated/numpy.ndindex.html

<!-- grade_id: cell-e5bedee04d14e5a8, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-de6936cc3a8d3be9, points: 5, grade: true, solution: false -->
```python
# Test case 1.
variables = ['x_1', 'x_2']
degrees = [1, 1]
coefficients1 = {(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}
coefficients2 = {(0, 0): -1, (1, 0): 1, (0, 1): 1, (1, 1): -1}
poly1 = MultivariatePolynomial(variables, degrees, coefficients1)
poly2 = MultivariatePolynomial(variables, degrees, coefficients2)
poly_sum = poly1.add(poly2)
assert all(poly_sum.get_coefficient(deg) == 0 for deg in [(0, 0), (1, 0), (0, 1), (1, 1)])
```

<!-- grade_id: cell-34a7d6ffb4302933, points: 5, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-49c7c5545f998907, points: 5, grade: true, solution: false -->
```python
# Hidden test case 2.
```

<!-- grade_id: cell-59c242bbfd28caad, points: 5, grade: true, solution: false -->
```python
# Hidden test case 3.
```

# 4. Polynomial Differentiation (25 points)

Implement a method `differentiate(self, var_index)` that takes the index of a variable and returns a new `MultivariatePolynomial` object representing the partial derivative of the polynomial with respect to that variable.

The mathematical expression for the differentiation is:

$$\frac{\partial P}{\partial x_j} = \sum_{i_1 = 0}^{d_1} \cdots \sum_{i_j = 1}^{d_j} \cdots \sum_{i_n = 0}^{d_n} i_j p_{i_1, \dots, i_n} x_1^{i_1} \cdots x_j^{i_j - 1} \cdots x_n^{i_n}$$

Return `None` if the variable index is invalid (not in the range of the number of variables).

Hint: You may use `numpy.ndindex` to iterate over all possible degree combinations. For more help please visit https://numpy.org/doc/2.3/reference/generated/numpy.ndindex.html

<!-- grade_id: cell-2511ccff781e55c7, points: 5, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-2eb0f117125c4a97, points: 5, grade: true, solution: false -->
```python
# Test case 1.
variables = ['x_1', 'x_2']
degrees = [1, 1]
coefficients = {(0, 0): 1, (1, 0): -1, (0, 1): -1, (1, 1): 1}
poly = MultivariatePolynomial(variables, degrees, coefficients)
poly_deriv = poly.differentiate(0)
expected = {
    (0, 0): -1,
    (0, 1): 1,
    (1, 0): None,
    (1, 1): None,
}
assert all(poly_deriv.get_coefficient(deg) == expected[deg] for deg in expected.keys())
```

<!-- grade_id: cell-7237ffda1555274a, points: 5, grade: true, solution: false -->
```python
# Hidden test case 1.
```

<!-- grade_id: cell-1225f7485a9b9759, points: 5, grade: true, solution: false -->
```python
# Hidden test case 2.
```

<!-- grade_id: cell-c60482398360f602, points: 5, grade: true, solution: false -->
```python
# Hidden test case 3.
```